Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

Do not edit or insert code into the test cells as all you insert will be overwritten by the automated testing code.

---

**Question 2**: Consider a Fourier series approximation of a function with the Fourier series given as $$S(x,n) = \frac{4}{\pi} \sum_{i=1}^{n}\,\frac{1}{2\,i-1}\,\sin\left(\frac{(2\,i-1)\,2\,\pi\,x}{T}\right)$$ for a fixed range of values $x/T$, i.e. values in the interval `]0,1[`, with $T$ the function period and $n$ the order of approximation, i.e. how many terms in the Fourier series should be used.

**Task**: Use `np.linspace(from, to, steps)` in order to create the $x/T$ values and plot all Fourier Series curves for the order parameters $n=(3, 10, 30)$ with appropriate axis labels and a title on a single figure. Also include a figure legend, describing your three curves. Observe how the Fourier series approximates the original step function better and better.

Therefore, **(a)** write a function `fourier(values, order)` which takes the x values from above as first input and the order parameter as a second input. It should calculate the corresponding Fourier sum and return the array of Fourier values.

Then **(b)** write a main program which creates the x values as described above and loops through the order parameters from above in order to plot the requested curves.

<!-- grade_id: Q2Code, points: 0, grade: false, solution: true -->
```python
# function first, then main script for plotting
# YOUR CODE HERE
```

<!-- grade_id: Q2testA, points: 2, grade: true, solution: false -->
```python
# Test cell
xt = np.array([0.01,0.25, 0.6])
# order test
assert (fourier(xt, 3)[0]<fourier(xt, 10)[0])
assert (fourier(xt, 3)[1]>fourier(xt, 10)[1])
```

<!-- grade_id: Q2testB, points: 1, grade: true, solution: false -->
```python
# Test cell
assert (fourier(xt, 3)[0]>0.22 and fourier(xt, 3)[0]<0.25)
```

### Tests - 3 marks; Style - 1 mark; Plot - 3 marks